# General tasks and directions

- Add your name, today's date, and the assignment title to the designated cell.
- Write your answers in the cells that contain `Add your answer here.` line.
- Write your code in the cells that contain `# Add your implementation here.` line.
- Use autograder tests that are provided for your convenience.
- Don't change or delete any provided code (including [cell magic](https://ipython.readthedocs.io/en/stable/interactive/magics.html) such as `%%capture output`).


<!-- grade_id: intro-author, points: 0, grade: true, solution: true -->
Add your answer here.

Name: Muhammad Mughees Akbar Warraich
Date: 14 March 2023

# Airline tweets

*Using `pandas` to clean data*.

This assignment is individual and you agree to submit your own work.


```python
import numpy as np
import pandas as pd

total_points = 0
task_points = 5
```

## Task 1

Read airline sentiment analysis tweets from the provided file *airline_tweets.csv*.

<!-- grade_id: solution-task-1, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df = pd.read_csv("airline_tweets.csv")

```

```python
# This cell is provided for your convenience to view the initial data status
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 14640 entries, 0 to 14639
Data columns (total 15 columns):
 #   Column                        Non-Null Count  Dtype  
---  ------                        --------------  -----  
 0   tweet_id                      14640 non-null  int64  
 1   airline_sentiment             14640 non-null  object 
 2   airline_sentiment_confidence  14640 non-null  float64
 3   negativereason                9178 non-null   object 
 4   negativereason_confidence     10522 non-null  float64
 5   airline                       14640 non-null  object 
 6   airline_sentiment_gold        40 non-null     object 
 7   name                          14640 non-null  object 
 8   negativereason_gold           32 non-null     object 
 9   retweet_count                 14640 non-null  int64  
 10  text                          14640 non-null  object 
 11  tweet_coord                   1019 non-null   object 
 12  tweet_created                 14640 non-null  object 
 13  t
```

```python
df
```

Output:
```
tweet_id airline_sentiment  airline_sentiment_confidence  \
0      570306133677760513           neutral                        1.0000   
1      570301130888122368          positive                        0.3486   
2      570301083672813571           neutral                        0.6837   
3      570301031407624196          negative                        1.0000   
4      570300817074462722          negative                        1.0000   
...                   ...               ...                           ...   
14635  569587686496825344          positive                        0.3487   
14636  569587371693355008          negative                        1.0000   
14637  569587242672398336           neutral                        1.0000   
14638  569587188687634433          negative                        1.0000   
14639  569587140490866689           neutral                        0.6771   

               negativereason  negativereason_confidence         airline  \
0               
```

<!-- grade_id: test-task-1-1, points: 5, grade: true, solution: false -->
```python
assert df.shape == (14640, 15)
total_points += task_points
```

<!-- grade_id: test-task-1-2, points: 5, grade: true, solution: false -->
```python
assert list(df.columns) == ['tweet_id', 'airline_sentiment', 'airline_sentiment_confidence',
                           'negativereason', 'negativereason_confidence', 'airline',
                           'airline_sentiment_gold', 'name', 'negativereason_gold',
                           'retweet_count', 'text', 'tweet_coord', 'tweet_created',
                           'tweet_location', 'user_timezone']
total_points += task_points
```

## Task 2

Remove columns where more than one third (`33%`) of the values are missing.
You must do that programmatically, not hard-coding the column names.

<!-- grade_id: solution-task-2, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df.dropna(thresh=9759,axis="columns",inplace=True)
```

<!-- grade_id: test-task-2-1, points: 5, grade: true, solution: false -->
```python
assert df.shape == (14640, 11)
total_points += task_points
```

<!-- grade_id: test-task-2-2, points: 5, grade: true, solution: false -->
```python
assert list(df.columns) == ['tweet_id', 'airline_sentiment', 'airline_sentiment_confidence',
                           'negativereason_confidence', 'airline', 'name', 'retweet_count',
                            'text', 'tweet_created', 'tweet_location', 'user_timezone']
total_points += task_points
```

## Task 3

Remove duplicate rows and rows with more than `2` missing values.

<!-- grade_id: solution-task-3, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df.drop_duplicates(inplace= True)
```

```python
df.dropna(thresh=9,axis=0,inplace=True)
```

```python
df.shape
```

Output:
```
(13980, 11)
```

<!-- grade_id: test-task-3-1, points: 5, grade: true, solution: false -->
```python
assert df.shape == (13980, 11)
total_points += task_points
```

<!-- grade_id: test-task-3-2, points: 5, grade: true, solution: false -->
```python
assert not (df.isna().sum(axis=1) > 2).any()
total_points += task_points
```

## Task 4

Impute missing values in the columns as follows:

- negativereason_confidence: the median value of the column
- tweet_location: "Decorah, IA"
- user_timezone: the most common value of the column


<!-- grade_id: solution-task-4, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df.tweet_location.fillna("Decorah, IA",inplace=True)

```

```python
median_value = df["negativereason_confidence"].median()
df["negativereason_confidence"].fillna(median_value, inplace=True)

```

```python
common = df['user_timezone'].value_counts().idxmax()
```

```python
common
```

Output:
```
'Eastern Time (US & Canada)'
```

```python
df.user_timezone.fillna(common,inplace=True)
```

<!-- grade_id: test-task-4-1, points: 5, grade: true, solution: false -->
```python
assert abs(df.negativereason_confidence.mean() - 0.645769) < 0.001
total_points += task_points
```

<!-- grade_id: test-task-4-2, points: 5, grade: true, solution: false -->
```python
assert f"{df.tweet_location.mode().iloc[0]} ({df.user_timezone.mode().iloc[0]})" == \
                "Decorah, IA (Eastern Time (US & Canada))"
total_points += task_points
```

## Task 5

Set the data types of the columns as follows:

| column                       | data type                      |
|------------------------------|--------------------------------|
| tweet_id                     | built-in 64-bit integer        |
| airline_sentiment            | `pandas` categorical           |
| airline_sentiment_confidence | built-in 64-bit floating-point |
| negativereason_confidence    | built-in 64-bit floating-point |
| airline                      | `pandas` string                |
| name                         | `pandas` string                |
| retweet_count                | built-in 64-bit integer        |
| text                         | `pandas` string                |
| tweet_created                | `pandas` datetime (UTC)        |
| tweet_location               | `pandas` string                |
| user_timezone                | `pandas` categorical           |



<!-- grade_id: solution-task-5, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df = df.astype({"tweet_id": int, "airline_sentiment": ('category'),"airline_sentiment_confidence": float})
df = df.astype({"negativereason_confidence": float,"airline": ('string'), "name": ('string'),"retweet_count": int})
df = df.astype({"text": ('string'), "tweet_created":('datetime64[ns, UTC]') ,"tweet_location": ('string'), "user_timezone":('category')})
















```

```python
df.dtypes
```

Output:
```
tweet_id                                      int64
airline_sentiment                          category
airline_sentiment_confidence                float64
negativereason_confidence                   float64
airline                                      string
name                                         string
retweet_count                                 int64
text                                         string
tweet_created                   datetime64[ns, UTC]
tweet_location                               string
user_timezone                              category
dtype: object
```

<!-- grade_id: test-task-5-1, points: 5, grade: true, solution: false -->
```python
assert df.shape == (13980, 11)
total_points += task_points
```

<!-- grade_id: test-task-5-2, points: 5, grade: true, solution: false -->
```python
assert list(df.dtypes) == ["int64", "category", "float64", "float64",
                           "string", "string", "int64", "string",
                           "datetime64[ns, UTC]", "string", "category"]
total_points += task_points
```

```python
# This cell is provided for your convenience to view the result
# All 11 columns must have 13980 non-null values and proper data types
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 13980 entries, 0 to 14639
Data columns (total 11 columns):
 #   Column                        Non-Null Count  Dtype              
---  ------                        --------------  -----              
 0   tweet_id                      13980 non-null  int64              
 1   airline_sentiment             13980 non-null  category           
 2   airline_sentiment_confidence  13980 non-null  float64            
 3   negativereason_confidence     13980 non-null  float64            
 4   airline                       13980 non-null  string             
 5   name                          13980 non-null  string             
 6   retweet_count                 13980 non-null  int64              
 7   text                          13980 non-null  string             
 8   tweet_created                 13980 non-null  datetime64[ns, UTC]
 9   tweet_location                13980 non-null  string             
 10  user_timezone                 1398
```

```python
print(f"Total points: {total_points}/50.")
```

Output:
```
Total points: 50/50.
```

## Submission Checklist

- [ ] Your name, today's date, and the assignment title in the designated cell.
- [ ] Your answers in the designated cells (if required).
- [ ] Your code runs and produces the expected output.
- [ ] The validity of your code is verified by autograders (if provided).
- [ ] Restart the kernel and run all cells (in the menubar, select *Kernel*, then *Restart Kernel and Run All Cells*).
- [ ] Save the notebook.
- [ ] Submit the assignment.


```python

```