# General tasks and directions

- Add your name, today's date, and the assignment title to the designated cell.
- Write your answers in the cells that contain `Add your answer here.` line.
- Write your code in the cells that contain `# Add your implementation here.` line.
- Use autograder tests that are provided for your convenience.
- Don't change or delete any provided code (including [cell magic](https://ipython.readthedocs.io/en/stable/interactive/magics.html) such as `%%capture output`).


## Add your name, today's date, and the assignment title

author: Muhammad Mughees Akbar Warraich

date: March 1 2023

assignment: Exercise 2


Resources used: 
    
#https://www.statology.org/or-operator-in-pandas/ (or operator)
#https://www.geeksforgeeks.org/how-to-reverse-the-column-order-of-the-pandas-dataframe/ (how to reverst a column in pandas)
#https://www.skytowner.com/explore/getting_shortest_and_longest_strings_in_pandas_dataframe (link for resource to find the shortest and longest string in a column)
#https://pandas.pydata.org/docs/reference/api/pandas.Series.str.startswith.html (starts with function)


# Fortune 1000 (2022)

```python
%config Completer.use_jedi = False

import numpy as np
import pandas as pd

from pandas import Series, DataFrame

total_points = 0
```

```python
df = pd.read_csv("fortune1000.csv")
df

```

Output:
```
company  rank  rank_change   revenue   profit  \
0               Walmart     1          0.0  572754.0  13673.0   
1                Amazon     2          0.0  469822.0  33364.0   
2                 Apple     3          0.0  365817.0  94680.0   
3            CVS Health     4          0.0  292111.0   7910.0   
4    UnitedHealth Group     5          0.0  287597.0  17285.0   
..                  ...   ...          ...       ...      ...   
995       Vizio Holding   996          0.0    2124.0    -39.4   
996   1-800-Flowers.com   997          0.0    2122.2    118.7   
997               Cowen   998          0.0    2112.8    295.6   
998             Ashland   999          0.0    2111.0    220.0   
999            DocuSign  1000          0.0    2107.2    -70.0   

     num. of employees       sector           city state newcomer ceo_founder  \
0            2300000.0    Retailing    Bentonville    AR       no          no   
1            1608000.0    Retailing        Seattle    WA       no        
```

```python
assert _.shape == (1000, 18)
```

## Number of employees at Kroger

<!-- grade_id: solution-task-0, points: 0, grade: true, solution: true -->


**Task 0**

You should work with the same `DataFrame` throughout the whole exercise but may add new calculted columns, if needed.
Assertions rely on the result of the previous cell execution (`_`), so you don't have to print anything but rather make sure the result of all operations in your solution matches the expected.

Try not to mutate `df` and restart the notebook from the top if necessary.

<!-- grade_id: solution-task-1, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[df.company == "Kroger"]["num. of employees"].iloc[0]
```

Output:
```
420000.0
```

<!-- grade_id: test-task-1, points: 2, grade: true, solution: false -->
```python
assert _ == 420000
total_points += 2
```

## Total market capitalization of all companies

<!-- grade_id: solution-task-2, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df['Market Cap']=pd.to_numeric(df['Market Cap'], errors='coerce')
df['Market Cap'].sum()

```

Output:
```
38462044.5
```

<!-- grade_id: test-task-2, points: 2, grade: true, solution: false -->
```python
assert _ == 38462044.5
total_points += 2
```

## Total number of people employed by all companies

<!-- grade_id: solution-task-3, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df['num. of employees'].sum()

```

Output:
```
35752880.0
```

<!-- grade_id: test-task-3, points: 2, grade: true, solution: false -->
```python
assert _ == 35752880.0
total_points += 2
```

## Number of unique sectors

<!-- grade_id: solution-task-4, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df.sector.nunique()

```

Output:
```
21
```

<!-- grade_id: test-task-4, points: 2, grade: true, solution: false -->
```python
assert _ == 21
total_points += 2
```

## All unique cities in California

<!-- grade_id: solution-task-5, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[df.state == "CA"]["city"].unique()

```

Output:
```
array(['Cupertino', 'Mountain View', 'San Ramon', 'Menlo Park',
       'San Francisco', 'Santa Clara', 'Burbank', 'Palo Alto', 'San Jose',
       'San Diego', 'Fremont', 'Los Gatos', 'Long Beach', 'Foster City',
       'Thousand Oaks', 'Dublin', 'Rosemead', 'Los Angeles',
       'Woodland Hills', 'Newport Beach', 'Santa Ana', 'Santa Monica',
       'San Mateo', 'Glendale', 'El Segundo', 'Oakland', 'Milpitas',
       'Redwood City', 'Manhattan Beach', 'Beverly Hills', 'Sunnyvale',
       'Corona', 'Irvine', 'Pleasanton', 'Santa Rosa', 'Sylmar',
       'Westlake Village', 'San Rafael', 'Foothill Ranch', 'Corte Madera',
       'Watsonville', 'Walnut Creek', 'Pasadena', 'Carlsbad',
       'Emeryville', 'Calabasas Hills', 'Hercules', 'Brea',
       'San Juan Capistrano', 'Goleta'], dtype=object)
```

<!-- grade_id: test-task-5, points: 2, grade: true, solution: false -->
```python
assert len(list(_)) == 50
total_points += 2
```

## All companies from Iowa

<!-- grade_id: solution-task-6, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[df.state == "IA"]
```

Output:
```
company  rank  rank_change  revenue  \
255                      Principal Financial   256        -50.0  14262.7   
444                   Casey's General Stores   445        -74.0   7654.2   
726  American Equity Investment Life Holding   727          0.0   3689.5   
788                   Renewable Energy Group   789          0.0   3244.1   
982                                      HNI   983          0.0   2184.4   

     profit  num. of employees              sector             city state  \
255  1710.6            18600.0          Financials       Des Moines    IA   
444   312.9            27583.0           Retailing           Ankeny    IA   
726   474.0              850.0          Financials  West Des Moines    IA   
788   213.8             1196.0              Energy             Ames    IA   
982    59.8             8000.0  Household Products        Muscatine    IA   

    newcomer ceo_founder ceo_woman profitable prev_rank                  CEO  \
255       no          no        no   
```

<!-- grade_id: test-task-6, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Principal Financial",
                              "Casey's General Stores",
                              "American Equity Investment Life Holding",
                              "Renewable Energy Group",
                              "HNI"]
total_points += 2
```

## All companies from Illinois in the Telecommunications sector

<!-- grade_id: solution-task-7, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df[(df.state == "IL") & (df.sector == "Telecommunications")]

```

Output:
```
company  rank  rank_change  revenue  profit  \
572  Telephone & Data Systems   573          0.0   5329.0   156.0   

     num. of employees              sector     city state newcomer  \
572             8800.0  Telecommunications  Chicago    IL       no   

    ceo_founder ceo_woman profitable prev_rank                   CEO  \
572          no        no        yes            LeRoy T. Carlson Jr.   

                    Website Ticker  Market Cap  
572  https://www.tdsinc.com    TDS      2163.4
```

<!-- grade_id: test-task-7, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Telephone & Data Systems"]
total_points += 2
```

## All newcomer companies with female CEO

<!-- grade_id: solution-task-8, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[(df.newcomer == "yes") & (df.ceo_woman == "yes")]

```

Output:
```
company  rank  rank_change  revenue  profit  num. of employees  \
394  Jackson Financial   395          0.0   8848.0  3183.0             3150.0   

         sector     city state newcomer ceo_founder ceo_woman profitable  \
394  Financials  Lansing    MI      yes          no       yes        yes   

    prev_rank                 CEO                  Website Ticker  Market Cap  
394            Laura L. Prieskorn  https://www.jackson.com    JXN      3895.1
```

<!-- grade_id: test-task-8, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Jackson Financial"]
total_points += 2
```

## All profitable newcomers from Ohio or Oklahoma

<!-- grade_id: solution-task-9, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[(df.newcomer == "yes") & (df.profitable == "yes") & ((df.state == "OH")|(df.state == "OK"))]

```

Output:
```
company  rank  rank_change  revenue  profit  \
170          Cleveland-Cliffs   171          0.0  20444.0  2988.0   
308              Devon Energy   309          0.0  12206.0  2813.0   
460  TravelCenters of America   461          0.0   7336.8    58.5   
479         Victoria's Secret   480          0.0   6784.6   646.4   

     num. of employees     sector           city state newcomer ceo_founder  \
170            26000.0  Materials      Cleveland    OH      yes          no   
308             1600.0     Energy  Oklahoma City    OK      yes          no   
460            16540.0  Retailing       Westlake    OH      yes          no   
479            24000.0  Retailing   Reynoldsburg    OH      yes          no   

    ceo_woman profitable prev_rank                    CEO  \
170        no        yes            C. Lourenco Goncalves   
308        no        yes              Richard E. Muncrief   
460        no        yes             Jonathan M. Pertchik   
479        no        yes            
```

<!-- grade_id: test-task-9, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Cleveland-Cliffs", 
                              "Devon Energy", 
                              "TravelCenters of America", 
                              "Victoria's Secret"]
total_points += 2
```

## Names of CEOs of all companies from Alabama

<!-- grade_id: solution-task-10, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[df.state == "AL"]
```

Output:
```
company  rank  rank_change  revenue  profit  num. of employees  \
488  Regions Financial   489        -55.0   6605.0  2521.0            19626.0   
561   Vulcan Materials   562          0.0   5552.2   670.8            11912.0   
591   Encompass Health   592          0.0   5121.6   412.2            35681.0   

          sector        city state newcomer ceo_founder ceo_woman profitable  \
488   Financials  Birmingham    AL       no          no        no        yes   
561    Materials  Birmingham    AL       no          no        no        yes   
591  Health Care  Birmingham    AL       no          no        no        yes   

    prev_rank                 CEO                          Website Ticker  \
488     434.0  John M. Turner Jr.          https://www.regions.com     RF   
561                J. Thomas Hill  https://www.vulcanmaterials.com    VMC   
591                  Mark J. Tarr  https://www.encompasshealth.com    EHC   

     Market Cap  
488     20860.9  
561     24412.7  
591   
```

<!-- grade_id: test-task-10, points: 2, grade: true, solution: false -->
```python
assert list(_["CEO"]) == ["John M. Turner Jr.", "J. Thomas Hill", "Mark J. Tarr"]
total_points += 2
```

## Names of founder CEOs of all small (fewer than 10,000 employees) Technology companies not from California

<!-- grade_id: solution-task-11, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df[(df.state != "CA") & (df['num. of employees'] < 10000) & (df.ceo_founder == "yes") & (df.sector == "Technology")]




```

Output:
```
company  rank  rank_change  revenue  profit  \
423         Zillow Group   424          0.0   8147.4  -527.8   
494              Compass   495          0.0   6421.0  -494.1   
757  Akamai Technologies   758          0.0   3461.2   651.6   

     num. of employees      sector       city state newcomer ceo_founder  \
423             8005.0  Technology    Seattle    WA      yes         yes   
494             4775.0  Technology   New York    NY      yes         yes   
757             8700.0  Technology  Cambridge    MA       no         yes   

    ceo_woman profitable prev_rank                  CEO  \
423        no         no              Richard N. Barton   
494        no         no              Robert L. Reffkin   
757        no        yes            F. Thomson Leighton   

                         Website Ticker  Market Cap  
423  https://www.zillowgroup.com     ZG     12287.1  
494      https://www.compass.com   COMP      3358.8  
757       https://www.akamai.com   AKAM     19210.0
```

<!-- grade_id: test-task-11, points: 2, grade: true, solution: false -->
```python
assert list(_["CEO"]) == ["Richard N. Barton", "Robert L. Reffkin", "F. Thomson Leighton"]
total_points += 2
```

## All companies with the market capitalization of $1,000,000,000,000 or more

Note that market cap in the data file is given in millions of dollars.

<!-- grade_id: solution-task-12, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df['Market Cap']=pd.to_numeric(df['Market Cap'], errors='coerce')
df[df['Market Cap'] >= 1000000]


```

Output:
```
company  rank  rank_change   revenue   profit  num. of employees  \
1      Amazon     2          0.0  469822.0  33364.0          1608000.0   
2       Apple     3          0.0  365817.0  94680.0           154000.0   
7    Alphabet     8          1.0  257637.0  76033.0           156500.0   
13  Microsoft    14          1.0  168088.0  61271.0           181000.0   

        sector           city state newcomer ceo_founder ceo_woman profitable  \
1    Retailing        Seattle    WA       no          no        no        yes   
2   Technology      Cupertino    CA       no          no        no        yes   
7   Technology  Mountain View    CA       no          no        no        yes   
13  Technology        Redmond    WA       no          no        no        yes   

   prev_rank              CEO              Website Ticker  Market Cap  
1        2.0  Andrew R. Jassy       www.amazon.com   AMZN   1202717.0  
2        3.0  Timothy D. Cook        www.apple.com   AAPL   2443962.0  
7        9.0 
```

<!-- grade_id: test-task-12, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Amazon", "Apple", "Alphabet", "Microsoft"]
total_points += 2
```

## A company that climbed the highest in the ranking

<!-- grade_id: solution-task-13, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df[df['rank_change']== df['rank_change'].max()]

```

Output:
```
company  rank  rank_change  revenue  profit  num. of employees  \
158  Blackstone   159        295.0  22577.1  5857.4             3795.0   

         sector      city state newcomer ceo_founder ceo_woman profitable  \
158  Financials  New York    NY       no         yes        no        yes   

    prev_rank                    CEO                     Website Ticker  \
158     454.0  Stephen A. Schwarzman  https://www.blackstone.com     BX   

     Market Cap  
158     88907.2
```

<!-- grade_id: test-task-13, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Blackstone"]
total_points += 2
```

## 3 profitable companies that dropped in the ranking the most

<!-- grade_id: solution-task-14, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df[df.profitable == "yes"].sort_values("rank_change", ascending=False).tail(3)[::-1]

```

Output:
```
company  rank  rank_change  revenue  profit  num. of employees  \
434  Bath & Body Works   435       -178.0   7881.5  1333.5            32850.0   
379           VF Corp.   380       -117.0   9238.8   407.9            33600.0   
337             Biogen   338       -110.0  10981.7  1556.1             9610.0   

          sector       city state newcomer ceo_founder ceo_woman profitable  \
434    Retailing   Columbus    OH       no          no       yes        yes   
379      Apparel     Denver    CO       no          no        no        yes   
337  Health Care  Cambridge    MA       no          no        no        yes   

    prev_rank               CEO                 Website Ticker  Market Cap  
434     257.0     Sarah E. Nash  https://www.bbwinc.com   BBWI     11419.9  
379     263.0     Steven Rendle     https://www.vfc.com    VFC     11620.0  
337     228.0  Michel Vounatsos  https://www.biogen.com   BIIB     30950.4
```

<!-- grade_id: test-task-14, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Bath & Body Works", "VF Corp.", "Biogen"]
total_points += 2
```

## All companies with fewer than 500 employees, sorted by their market cap from highest to lowest

<!-- grade_id: solution-task-15, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df['Market Cap']=pd.to_numeric(df['Market Cap'], errors='coerce')
```

<!-- grade_id: solution-task-15, points: 0, grade: false, solution: true -->
```python
df[(df['num. of employees'] < 500)].sort_values("Market Cap", ascending=False)
```

Output:
```
company  rank  rank_change  revenue  profit  \
629                  Welltower   630          0.0   4742.1   336.1   
710                     Ventas   711          0.0   3828.0    49.0   
855      Host Hotels & Resorts   856          0.0   2890.0   -11.0   
866  Annaly Capital Management   867          0.0   2836.2  2389.9   
445     A-Mark Precious Metals   446         49.0   7613.0   159.6   

     num. of employees       sector        city state newcomer ceo_founder  \
629              464.0   Financials      Toledo    OH       no          no   
710              434.0   Financials     Chicago    IL       no          no   
855              160.0   Financials    Bethesda    MD       no          no   
866              171.0   Financials    New York    NY       no          no   
445              347.0  Wholesalers  El Segundo    CA       no          no   

    ceo_woman profitable prev_rank                   CEO  \
629        no        yes                 Shankh S. Mitra   
710       yes
```

<!-- grade_id: test-task-15, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Welltower", 
                              "Ventas", 
                              "Host Hotels & Resorts", 
                              "Annaly Capital Management", 
                              "A-Mark Precious Metals"]
total_points += 2
```

## Top 5 companies with the highest profit

<!-- grade_id: solution-task-16, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df.sort_values("profit", ascending=False).head(5)

```

Output:
```
company  rank  rank_change   revenue   profit  \
2                Apple     3          0.0  365817.0  94680.0   
6   Berkshire Hathaway     7         -1.0  276094.0  89795.0   
7             Alphabet     8          1.0  257637.0  76033.0   
13           Microsoft    14          1.0  168088.0  61271.0   
23      JPMorgan Chase    24         -5.0  127202.0  48334.0   

    num. of employees      sector           city state newcomer ceo_founder  \
2            154000.0  Technology      Cupertino    CA       no          no   
6            372000.0  Financials          Omaha    NE       no          no   
7            156500.0  Technology  Mountain View    CA       no          no   
13           181000.0  Technology        Redmond    WA       no          no   
23           271025.0  Financials       New York    NY       no          no   

   ceo_woman profitable prev_rank                CEO  \
2         no        yes       3.0    Timothy D. Cook   
6         no        yes       6.0  Warren E
```

<!-- grade_id: test-task-16, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Apple", "Berkshire Hathaway", "Alphabet", "Microsoft", "JPMorgan Chase"]
total_points += 2
```

## Bottom 5 companies by profit (i.e. highest losses)

<!-- grade_id: solution-task-17, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df.sort_values("profit", ascending=True).head()[::-1]
```

Output:
```
company  rank  rank_change   revenue  profit  num. of employees  \
232   Cheniere Energy   233         95.0   15864.0 -2343.0             1550.0   
59             Boeing    60         -6.0   62286.0 -4202.0           142000.0   
913            WeWork   914          0.0    2570.1 -4439.0             4400.0   
8            McKesson     9         -2.0  238228.0 -4539.0            67500.0   
47   General Electric    48        -10.0   74196.0 -6520.0           168000.0   

                  sector      city state newcomer ceo_founder ceo_woman  \
232               Energy   Houston    TX       no          no        no   
59   Aerospace & Defense   Chicago    IL       no          no        no   
913           Financials  New York    NY       no          no        no   
8            Health Care    Irving    TX       no          no        no   
47           Industrials    Boston    MA       no          no        no   

    profitable prev_rank                   CEO                   Website  \

```

<!-- grade_id: test-task-17, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Cheniere Energy", "Boeing", "WeWork", "McKesson", "General Electric"]
total_points += 2
```

## Top 5 companies with the highest market capitalization

<!-- grade_id: solution-task-18, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.


df['Market Cap']=pd.to_numeric(df['Market Cap'], errors='coerce')
df.sort_values('Market Cap', ascending=False).head(5)


```

Output:
```
company  rank  rank_change   revenue   profit  num. of employees  \
2       Apple     3          0.0  365817.0  94680.0           154000.0   
13  Microsoft    14          1.0  168088.0  61271.0           181000.0   
7    Alphabet     8          1.0  257637.0  76033.0           156500.0   
1      Amazon     2          0.0  469822.0  33364.0          1608000.0   
64      Tesla    65         35.0   53823.0   5519.0            99290.0   

                    sector           city state newcomer ceo_founder  \
2               Technology      Cupertino    CA       no          no   
13              Technology        Redmond    WA       no          no   
7               Technology  Mountain View    CA       no          no   
1                Retailing        Seattle    WA       no          no   
64  Motor Vehicles & Parts         Austin    TX       no         yes   

   ceo_woman profitable prev_rank              CEO              Website  \
2         no        yes       3.0  Timothy D. Cook   
```

<!-- grade_id: test-task-18, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Apple", "Microsoft", "Alphabet", "Amazon", "Tesla"]
total_points += 2
```

## Top 5 companies with the highest profit per employee

<!-- grade_id: solution-task-19, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df["profit per employee"] = df["profit"]/ df["num. of employees"]
df.sort_values("profit per employee", ascending=False).head(5)

```

Output:
```
company  rank  rank_change   revenue   profit  \
866  Annaly Capital Management   867          0.0    2836.2   2389.9   
194                    Moderna   195          0.0   18471.0  12202.0   
32                  Fannie Mae    33         -8.0  101543.0  22176.0   
478         Diamondback Energy   479          0.0    6797.0   2182.0   
609                        TPG   610          0.0    4976.4   2180.9   

     num. of employees       sector        city state newcomer ceo_founder  \
866              171.0   Financials    New York    NY       no          no   
194             2700.0  Health Care   Cambridge    MA      yes          no   
32              7400.0   Financials  Washington    DC       no          no   
478              870.0       Energy     Midland    TX      yes          no   
609             1013.0   Financials  Fort Worth    TX       no          no   

    ceo_woman profitable prev_rank                   CEO  \
866        no        yes            David L. Finkelstein   
1
```

<!-- grade_id: test-task-19, points: 1, grade: true, solution: false -->
```python
assert [c for c in _["company"]] == ["Annaly Capital Management",
                                     "Moderna",
                                     "Fannie Mae",
                                     "Diamondback Energy",
                                     "TPG"]
total_points += 2
```

## All companies with a founder serving as a CEO and a 1-character ticker

<!-- grade_id: solution-task-20, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[df.Ticker.apply(lambda x: len(str(x)))==1 & (df.ceo_founder == "yes")]

```

Output:
```
company  rank  rank_change  revenue  profit  num. of employees  \
274  Wayfair   275        -58.0  13708.0  -131.0            16681.0   

        sector    city state newcomer ceo_founder ceo_woman profitable  \
274  Retailing  Boston    MA       no         yes        no         no   

    prev_rank            CEO                  Website Ticker  Market Cap  \
274     217.0  Niraj S. Shah  https://www.wayfair.com      W     11641.6   

     profit per employee  
274            -0.007853
```

<!-- grade_id: test-task-20, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Wayfair"]
total_points += 2
```

## All companies with a 5-character ticker

<!-- grade_id: solution-task-21, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
df[df.Ticker.apply(lambda x: len(str(x)))==5]

```

Output:
```
company  rank  rank_change   revenue   profit  \
7            Alphabet     8          1.0  257637.0  76033.0   
27            Comcast    28         -2.0  116385.0  14159.0   
264     Qurate Retail   265        -49.0   14044.0    340.0   
325     Liberty Media   326          0.0   11400.0    398.0   
589  Rush Enterprises   590          0.0    5126.1    241.4   
608    Ingles Markets   609          0.0    4987.9    249.7   
617    Kelly Services   618          0.0    4909.7    156.1   

     num. of employees              sector            city state newcomer  \
7             156500.0          Technology   Mountain View    CA       no   
27            189000.0  Telecommunications    Philadelphia    PA       no   
264            26745.0           Retailing       Englewood    CO       no   
325             6886.0               Media       Englewood    CO       no   
589             7141.0           Retailing   New Braunfels    TX       no   
608            18460.0  Food & Drug Stores  Bla
```

<!-- grade_id: test-task-21, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Alphabet",
                              "Comcast",
                              "Qurate Retail",
                              "Liberty Media",
                              "Rush Enterprises",
                              "Ingles Markets",
                              "Kelly Services"]
total_points += 2
```

## Number of unique cities

Keep in mind that there are many cities with the same name in two different states, and Burlington and Columbus appear in three states.

<!-- grade_id: solution-task-22, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df.groupby(['city','state']).size().size


```

Output:
```
419
```

<!-- grade_id: test-task-22, points: 2, grade: true, solution: false -->
```python
assert _ == 419
total_points += 2
```

## Companies with the shortest name


<!-- grade_id: solution-task-23, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.


s_length = df['company'].str.len()
bool_mask = (s_length == s_length.min())
df[bool_mask]


```

Output:
```
company  rank  rank_change  revenue  profit  num. of employees  \
58       HP    59         -3.0  63487.0  6503.0            51000.0   
101      3M   102         -6.0  35355.0  5921.0            95000.0   
719      RH   720          0.0   3758.8   688.5             6100.0   
907      F5   908          0.0   2603.4   331.2             6429.0   

         sector          city state newcomer ceo_founder ceo_woman profitable  \
58   Technology     Palo Alto    CA       no          no        no        yes   
101   Chemicals      St. Paul    MN       no          no        no        yes   
719   Retailing  Corte Madera    CA       no          no        no        yes   
907  Technology       Seattle    WA       no          no        no        yes   

    prev_rank                   CEO             Website Ticker  Market Cap  \
58       56.0      Enrique J. Lores          www.hp.com    HPQ     33951.0   
101      96.0      Michael F. Roman          www.3m.com    MMM     75190.0   
719          
```

<!-- grade_id: test-task-23, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["HP", "3M", "RH", "F5"]
total_points += 2
```

## Companies whose CEO(s) has(-ve) the longest combined name


<!-- grade_id: solution-task-24, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.



s_length = df['CEO'].str.len()
bool_mask = (s_length == s_length.max())
df[bool_mask]

```

Output:
```
company  rank  rank_change  revenue  profit  num. of employees  \
53   Energy Transfer    54         27.0  67417.0  5470.0            12558.0   
288           Markel   289         22.0  12846.4  2425.0            20300.0   

         sector        city state newcomer ceo_founder ceo_woman profitable  \
53       Energy      Dallas    TX       no          no        no        yes   
288  Financials  Glen Allen    VA       no          no        no        yes   

    prev_rank                                    CEO                 Website  \
53       81.0  Thomas E. Long/Marshall S. McCrea III  www.energytransfer.com   
288     311.0  Thomas S. Gayner/Richard R. Whitt III  https://www.markel.com   

    Ticker  Market Cap  profit per employee  
53      ET     32182.0             0.435579  
288    MKL     20037.5             0.119458
```

<!-- grade_id: test-task-24, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Energy Transfer", "Markel"]
total_points += 2
```

## All companies with a website URL that starts with `https://https://`

<!-- grade_id: solution-task-25, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

df[df['Website'].str.startswith('https://https://')]
```

Output:
```
company  rank  rank_change  revenue  profit  \
309    Warner Bros. Discovery   310        -20.0  12191.0  1006.0   
545              Under Armour   546          0.0   5683.5   360.1   
620                  DoorDash   621          0.0   4888.0  -468.0   
642                  Covetrus   643          0.0   4575.0   -54.0   
865  Univision Communciations   866          0.0   2841.0   133.7   
912                 Pinterest   913          0.0   2578.0   316.4   

     num. of employees       sector           city state newcomer ceo_founder  \
309            11000.0        Media       New York    NY       no          no   
545            12300.0      Apparel      Baltimore    MD       no          no   
620             8600.0   Technology  San Francisco    CA       no         yes   
642             4873.0  Health Care       Portland    ME       no          no   
865                NaN        Media       New York    NY       no          no   
912             3225.0   Technology  San Francisco  
```

<!-- grade_id: test-task-25, points: 2, grade: true, solution: false -->
```python
assert list(_["company"]) == ["Warner Bros. Discovery",
                              "Under Armour",
                              "DoorDash",
                              "Covetrus",
                              "Univision Communciations",
                              "Pinterest"]
total_points += 2
```

```python
print(f"Total points: {total_points}/50")
```

Output:
```
Total points: 50/50
```

## Submission Checklist

- [ ] Your name, today's date, and the assignment title in the designated cell.
- [ ] Your answers in the designated cells (if required).
- [ ] Your code runs and produces the expected output.
- [ ] The validity of your code is verified by autograders (if provided).
- [ ] Restart the kernel and run all cells (in the menubar, select *Kernel*, then *Restart Kernel and Run All Cells*).
- [ ] Save the notebook.
- [ ] Submit the assignment.


```python

```