# General tasks and directions

- Add your name, today's date, and the assignment title to the designated cell.
- Write your answers in the cells that contain `Add your answer here.` line.
- Write your code in the cells that contain `# Add your implementation here.` line.
- Use autograder tests that are provided for your convenience.
- Don't change or delete any provided code (including [cell magic](https://ipython.readthedocs.io/en/stable/interactive/magics.html) such as `%%capture output`).


## Add your name, today's date, and the assignment title

author: Muhammad Mughees Akbar Warraich

date: 02/19/2023

assignment: Project1



# Project 1

Working with `numpy`

This assignment is individual and you agree to submit your own work.


```python
import numpy as np

np.set_printoptions(precision=2, suppress=True, linewidth=120)
```

## Task 1

Read *housing.csv*. Ignore/delete the first row (headers) and the last column (*ocean_proximity*).
Store the remaining data as an array of floating-point numbers.


<!-- grade_id: solution-task-1, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

from numpy import genfromtxt
data = genfromtxt('housing.csv', delimiter = ',')

arr1 = np.delete(np.delete(data, 0, 0), 9, 1)

data = arr1.astype(np.float64)

```

<!-- grade_id: test-task-1-1, points: 5, grade: true, solution: false -->
```python
expected = (20640, 9)
assert data.shape == expected, f"Result is {data.shape} instead of the expected {expected}"
```

<!-- grade_id: test-task-1-2, points: 5, grade: true, solution: false -->
```python
assert data.dtype == "float64"
```

<!-- grade_id: test-task-1-3, points: 5, grade: true, solution: false -->
```python
expected = 25646.428856575825
assert abs(np.std(data[320]) - expected) < 0.001, \
    f"Result is {np.std(data[320])} instead of the expected {expected}"
```

<!-- grade_id: test-task-1-4, points: 5, grade: true, solution: false -->
```python
expected = 32378.774753352634
assert abs(np.std(data[1861]) - expected) < 0.001, \
    f"Result is {np.std(data[1861])} instead of the expected {expected}"
```

<!-- grade_id: test-task-1-5, points: 5, grade: true, solution: false -->
```python
expected = 13749.02957764602
assert abs(np.std(data[2023]) - expected) < 0.001, \
    f"Result is {np.std(data[2023])} instead of the expected {expected}"
```

## Task 2

Delete rows with missing values.


<!-- grade_id: solution-task-2, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.
 
rows_missing= np.isnan(data).any(1)
data =  data[~rows_missing] 
```

<!-- grade_id: test-task-2-1, points: 5, grade: true, solution: false -->
```python
expected = (20433, 9)
assert data.shape == expected, f"Result is {data.shape} instead of the expected {expected}"
```

<!-- grade_id: test-task-2-2, points: 5, grade: true, solution: false -->
```python
assert data.dtype == "float64"
```

<!-- grade_id: test-task-2-3, points: 5, grade: true, solution: false -->
```python
expected = 26857.35024532572
assert abs(np.std(data[320]) - expected) < 0.001, \
    f"Result is {np.std(data[320])} instead of the expected {expected}"
```

<!-- grade_id: test-task-2-4, points: 5, grade: true, solution: false -->
```python
expected = 62769.475045144216
assert abs(np.std(data[1861]) - expected) < 0.001, \
    f"Result is {np.std(data[1861])} instead of the expected {expected}"
```

<!-- grade_id: test-task-2-5, points: 5, grade: true, solution: false -->
```python
expected = 15614.362642183494
assert abs(np.std(data[2023]) - expected) < 0.001, \
    f"Result is {np.std(data[2023])} instead of the expected {expected}"
```

## Task 3

Add a column with a ratio of *total_bedrooms/households*.


<!-- grade_id: solution-task-3, points: 0, grade: false, solution: true -->
```python
# Add your implementation here.

arr = data[:, 4]
arr2 = data[:, 6]

ratio = arr/arr2

data = np.column_stack((data, ratio))

```

<!-- grade_id: test-task-3-1, points: 5, grade: true, solution: false -->
```python
expected = (20433, 10)
assert data.shape == expected, f"Result is {data.shape} instead of the expected {expected}"
```

<!-- grade_id: test-task-3-2, points: 5, grade: true, solution: false -->
```python
assert data.dtype == "float64"
```

<!-- grade_id: test-task-3-3, points: 5, grade: true, solution: false -->
```python
expected = 25649.75666055152
assert abs(np.std(data[320]) - expected) < 0.001, \
    f"Result is {np.std(data[320])} instead of the expected {expected}"
```

<!-- grade_id: test-task-3-4, points: 5, grade: true, solution: false -->
```python
expected = 59928.284864493144
assert abs(np.std(data[1861]) - expected) < 0.001, \
    f"Result is {np.std(data[1861])} instead of the expected {expected}"
```

<!-- grade_id: test-task-3-5, points: 5, grade: true, solution: false -->
```python
expected = 14916.751110074121
assert abs(np.std(data[2023]) - expected) < 0.001, \
    f"Result is {np.std(data[2023])} instead of the expected {expected}"
```

## Task 4

Calculate `min`, `max`, `mean`, `stdev`, `median` on all columns and store the results in a new array `housing_stats` as follows:
- the first row of the `housing_stats` are **minimum** values of each column
- the second row of the `housing_stats` are **maximum** values of each column
- the third row of the `housing_stats` are **average** values of each column
- the forth row of the `housing_stats` are **median** values of each column
- the fifth row of the `housing_stats` are **standard deviation** values of each column


*housing_stats.npy* is provided as a reference but you are not supposed to load the value

<!-- grade_id: solution-task-4, points: 0, grade: false, solution: true -->
```python

minimums = np.min(data, 0)
maximums = np.max(data, 0)
averages = np.mean(data, 0)
medians = np.median(data, 0)
stdDevs = np.std(data, 0)


housing_stats = np.array([minimums, maximums, averages, medians, stdDevs])

```

<!-- grade_id: test-task-4-1, points: 5, grade: true, solution: false -->
```python
print("Testing minimum values")
expected = [-124.35, 32.54, 1. , 2. , 1. , 3. , 1. , 0.5 , 14999. , 0.33]
assert np.all(abs(housing_stats[0] - expected < 0.01)), \
    f"Result is {housing_stats[0]} instead of the expected {expected}"
print("Done testing minimum values")
```

Output:
```
Testing minimum values
Done testing minimum values
```

<!-- grade_id: test-task-4-2, points: 5, grade: true, solution: false -->
```python
print("Testing maximum values")
expected = [-114.31, 41.95, 52. , 39320. , 6445. , 35682. , 6082. , 15. , 500001. , 34.07]
assert np.all(abs(housing_stats[1] - expected < 0.01)), \
    f"Result is {housing_stats[1]} instead of the expected {expected}"
print("Done testing maximum values")
```

Output:
```
Testing maximum values
Done testing maximum values
```

<!-- grade_id: test-task-4-3, points: 5, grade: true, solution: false -->
```python
print("Testing average values")
expected = [-119.57, 35.63, 28.63, 2636.5 , 537.87, 1424.95, 499.43, 3.87, 206864.41, 1.1 ]
assert np.all(abs(housing_stats[2] - expected < 0.01)), \
    f"Result is {housing_stats[2]} instead of the expected {expected}"
print("Done testing average values")
```

Output:
```
Testing average values
Done testing average values
```

<!-- grade_id: test-task-4-4, points: 5, grade: true, solution: false -->
```python
print("Testing median values")
expected =[-118.49, 34.26, 29. , 2127. , 435. , 1166. , 409. , 3.54, 179700. , 1.05]
assert np.all(abs(housing_stats[3] - expected < 0.01)), \
    f"Result is {housing_stats[3]} instead of the expected {expected}"
print("Done testing median values")
```

Output:
```
Testing median values
Done testing median values
```

<!-- grade_id: test-task-4-5, points: 5, grade: true, solution: false -->
```python
print("Testing standard deviation values")
expected = [2. , 2.14, 12.59, 2185.22, 421.37, 1133.18, 382.29, 1.9 , 115432.84 , 0.48]
assert np.all(abs(housing_stats[4] - expected < 0.01)), \
    f"Result is {housing_stats[4]} instead of the expected {expected}"
print("Done testing standard deviation values")
```

Output:
```
Testing standard deviation values
Done testing standard deviation values
```

```python
print("This is the final test to verify the whole solution")
assert np.array_equal(housing_stats, np.load("housing_stats.npy"))
print("Done!")
```

Output:
```
This is the final test to verify the whole solution
Done!
```

## Submission Checklist

- [ ] Your name, today's date, and the assignment title in the designated cell.
- [ ] Your answers in the designated cells (if required).
- [ ] Your code runs and produces the expected output.
- [ ] The validity of your code is verified by autograders (if provided).
- [ ] Restart the kernel and run all cells (in the menubar, select *Kernel*, then *Restart Kernel and Run All Cells*).
- [ ] Save the notebook.
- [ ] Submit the assignment.


```python

```