# General tasks and directions

- Add your name, today's date, and the assignment title to the designated cell.
- Write your answers in the cells that contain `Add your answer here.` line.
- Write your code in the cells that contain `# Add your implementation here.` line.
- Use autograder tests that are provided for your convenience.
- Don't change or delete any provided code (including [cell magic](https://ipython.readthedocs.io/en/stable/interactive/magics.html) such as `%%capture output`).


## Add your name, today's date, and the assignment title

author: Muhammad Mughees Akbar Warraich

date: 02/06/2023

assignment: 0


# Project 0

*Matrix multiplication*

This assignment is individual and you agree to submit your own work.


```python
# The next line enables autocomplete on Tab
%config Completer.use_jedi = False

# Some tasks are graded automatically
auto_assignment_points = 0

# Some tasks are graded manually
manual_assignment_points = 0

# Correct will appear in green
correct_answer_msg = "\x1b[32mCorrect.\x1b[0m"

# Error massage will appear in red
incorrect_answer_msg = "\x1b[31mSomething is wrong. Please check your code and/or the resulting output.\x1b[0m"

# We need pytest to test exceptions
import pytest
```

## Task

Define function `read_matrix` with a single parameter, string `filename`, that reads a matrix stored in a file and returns it as a list of lists.

Define function `mmult` with two parameters, matrices `m1` and `m2`, that returns the product of those two matrices.

[Matrix multiplication - Wikipedia](https://en.wikipedia.org/wiki/Matrix_multiplication)

<!-- grade_id: cell-e8d30513f149a5c7, points: 0, grade: false, solution: true -->
```python
def read_matrix(filename: str) -> list[list[int]]:
    list = []
    biglist = []
    file = open(filename, "r")
    for line in file:
        list = [int(num) for num in line.split()]
        biglist.append(list)
    return biglist




```

<!-- grade_id: cell-80df9674701a50c3, points: 0, grade: true, solution: false -->
```python
assert read_matrix("matrix1.txt") == [[1, 2, 3, 4], [5, 6, 7, 8], [9, 0, 1, 2]], incorrect_answer_msg
assert read_matrix("matrix2.txt") == [[1, 2], [3, 4], [5, 6], [7, 8]], incorrect_answer_ms

print(correct_answer_msg)
```

Output:
```
[32mCorrect.[0m
```

<!-- grade_id: cell-80df9674701a50c4, points: 0, grade: true, solution: false -->
```python
assert read_matrix("matrixA.txt") == [[1, 0, 1], [2, 1, 1], [0, 1, 1], [1, 1, 2]], incorrect_answer_msg
assert read_matrix("matrixB.txt") == [[1, 2, 1], [2, 3, 1], [4, 2, 2]], incorrect_answer_msg

print(correct_answer_msg)
```

Output:
```
[32mCorrect.[0m
```

<!-- grade_id: cell-c27f46d40bb2103f, points: 0, grade: false, solution: true -->
```python
def mmult(m1: list[list[int]], m2: list[list[int]]) -> list[list[int]]:
    """Multiply two matrices and return the result"""
    # Add your implementation here.
    a = len(m1)
    b = len(m1[0])
    c = len(m2)
    d = len(m2[0])
    total = []
    for row in range (a):
        currentrow = []
        for column in range (d):
            currentrow.append(0)
        total.append(currentrow)
        
    for x in range(a):
        for y in range(d):
            currentvalue = 0
            for k in range(b):
                currentvalue += m1[x][k]*m2[k][y]
            total[x][y] = currentvalue
    return total
    
    


   

```

<!-- grade_id: cell-5ccdaa77326d5495, points: 0, grade: true, solution: false -->
```python
m_1 = read_matrix("matrix1.txt")
m_2 = read_matrix("matrix2.txt")
assert mmult(m_1, m_2) == [[50, 60], [114, 140], [28, 40]]

print(correct_answer_msg)
```

Output:
```
[32mCorrect.[0m
```

<!-- grade_id: cell-58b2b988e54c6c5c, points: 0, grade: true, solution: false -->
```python
m_A = read_matrix("matrixA.txt")
m_B = read_matrix("matrixB.txt")
assert mmult(m_A, m_B) == [[5, 4, 3], [8, 9, 5], [6, 5, 3], [11, 9, 6]]

print(correct_answer_msg)
```

Output:
```
[32mCorrect.[0m
```

## Submission Checklist

- [ ] Your name, today's date, and the assignment title in the designated cell.
- [ ] Your answers in the designated cells (if required).
- [ ] Your code runs and produces the expected output.
- [ ] The validity of your code is verified by autograders (if provided).
- [ ] Restart the kernel and run all cells (in the menubar, select *Kernel*, then *Restart Kernel and Run All Cells*).
- [ ] Save the notebook.
- [ ] Submit the assignment.
